<!--NAVIGATION-->

<a href="https://colab.research.google.com/github/bpesquet/machine-learning-katas/blob/master/notebooks/katas/workflow/ClassificationMetrics.ipynb"><img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open in Colab" title="Open in Google Colaboratory"></a>


## Instructions

This is a self-correcting exercise generated by [nbgrader](https://github.com/jupyter/nbgrader). 

Fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Run subsequent cells to check your code.

---

# Kata: Code Classification Metrics From Scratch

## Package setup

```python
import numpy as np
```

## Accuracy

### Question

Complete the following function to compute accuracy.

<!-- grade_id: cell-6a21120c3e66ee57, points: 0, grade: false, solution: true -->
```python
def accuracy_score(y_true, y_pred):
    nb_valid = np.absolute(1-(np.absolute(y_true - y_pred)).sum())
    return nb_valid/len(y_pred)
```

<!-- grade_id: cell-29d49954b89dc997, points: 1, grade: true, solution: false -->
```python
y_true = np.array([2, 0, 2, 2, 0, 1])
y_pred = np.array([0, 1, 2, 2, 0, 2])

acc = accuracy_score(y_true, y_pred)
print(f'Accuracy: {acc * 100:.2f}%')

assert acc == 0.5
```

Output:
```
Accuracy: 50.00%
```

## TODO

- Add other metrics: precision, recall, f1-score
- Add confusion matrix

```python
y_true = np.array([0, 0, 0, 1, 1, 0, 1, 0, 1, 1])
y_pred = np.array([1, 1, 0, 0, 0, 1, 1, 0, 1, 1])
```

```python
def nb_true_positives(y_true, y_pred):
    nb_true_malign = np.array(np.where(y_true==1)[0])
    nb_pred_malign = np.array(np.where(y_pred==1)[0])
    nb_tp = [(np.where(nb_true_malign==i)[0]) for i in nb_pred_malign]
    nb_tp = len([i for i in nb_tp if len(i)!=0])
    return nb_tp
```

```python
def precision(y_true, y_pred):
    nb_tp = nb_true_positives(y_true, y_pred)
    nb_tp_fp = len(np.where(y_pred==1)[0])
    return nb_tp/nb_tp_fp
    
prec = precision(y_true, y_pred)
print(f'Precision: {prec * 100:.2f}%')

assert prec == 0.5
```

Output:
```
Precision: 50.00%
```

```python
def recall(y_true, y_pred):
    nb_tp = nb_true_positives(y_true, y_pred)
    nb_true_malign = len(np.array(np.where(y_true==1)[0]))
    return nb_tp/nb_true_malign

rec = recall(y_true, y_pred)
print(f'Recall: {rec * 100:.2f}')

assert rec == 0.6
```

Output:
```
Recall: 60.00
```

```python
def f1_score(y_true, y_pred):
    prec = precision(y_true, y_pred)
    rec = recall(y_true, y_pred)
    return 2*(prec*rec)/(prec+rec)
f1_score = f1_score(y_true, y_pred)
print(f'F1_score: {f1_score:.2f}')
```

Output:
```
F1_score: 0.55
```