<!--NAVIGATION-->

<a href="https://colab.research.google.com/github/bpesquet/machine-learning-katas/blob/master/notebooks/katas/algorithms/KNN_BreastCancer.ipynb"><img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open in Colab" title="Open in Google Colaboratory"></a>


## Instructions

This is a self-correcting exercise generated by [nbgrader](https://github.com/jupyter/nbgrader). 

Fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Run subsequent cells to check your code.

---

# Kata: Diagnose Breast Tumors with K-Nearest Neighbors

In this kata, you'll use a K-Nearest Neighbors classifier to help diagnose breast tumors.

The [Breast Cancer][1] dataset is used for multivariate binary classification between benign and maligant tumors. There are 569 total samples with 30 features each. Features were computed from a digitized image of a fine needle aspirate of a breast mass. They describe characteristics of the cell nuclei present in the image.

![](images/breast-cancer-logo.jpg)

[1]: https://archive.ics.uci.edu/ml/datasets/Breast+Cancer+Wisconsin+(Diagnostic)

## Package setup

```python
# Import base packages
import numpy as np
import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
import pandas as pd
```

```python
# Setup plots
%matplotlib inline
plt.rcParams['figure.figsize'] = 10, 8
%config InlineBackend.figure_format = 'retina'
sns.set()
```

### Question

Import the needed packages.

<!-- grade_id: cell-067da4adf1d7badc, points: 0, grade: false, solution: true -->
```python
# Import ML packages (edit this list if needed)
from sklearn.datasets import load_breast_cancer
from sklearn.datasets import make_classification
from sklearn.neighbors import KNeighborsClassifier
from sklearn.model_selection import train_test_split
from sklearn.preprocessing import StandardScaler
```

## Step 1: Loading the data

```python
dataset = load_breast_cancer()

# Put data in a pandas DataFrame
df_breast_cancer = pd.DataFrame(dataset.data, columns=dataset.feature_names)
# Add target and class to DataFrame
df_breast_cancer['target'] = dataset.target
df_breast_cancer['class'] = dataset.target_names[dataset.target]
# Show 10 random samples
df_breast_cancer.sample(n=10)
```

Output:
```
mean radius  mean texture  mean perimeter  mean area  mean smoothness  \
427       10.800         21.98           68.79      359.9          0.08801   
297       11.760         18.14           75.00      431.1          0.09968   
113       10.510         20.19           68.64      334.2          0.11220   
5         12.450         15.70           82.57      477.1          0.12780   
244       19.400         23.50          129.10     1155.0          0.10270   
452       12.000         28.23           76.77      442.5          0.08437   
283       16.240         18.77          108.80      805.1          0.10660   
504        9.268         12.87           61.49      248.7          0.16340   
267       13.590         21.84           87.16      561.0          0.07956   
507       11.060         17.12           71.25      366.5          0.11940   

     mean compactness  mean concavity  mean concave points  mean symmetry  \
427           0.05743         0.03614              0.01404         0.
```

## Step 2: Preparing the data

### Question

Compute the number of features of the dataset into the `num_features` variable.

<!-- grade_id: cell-aeab76633e8adca7, points: 0, grade: false, solution: true -->
```python
num_features = len(dataset.feature_names)
```

<!-- grade_id: cell-e6c24b765def5139, points: 1, grade: true, solution: false -->
```python
print(f'Number of features: {num_features}')

assert num_features == 30
```

Output:
```
Number of features: 30
```

### Question

In order to evaluate class distribution, compute the number of benign and malignant tumors into the `num_benign` and `num_malignant` variables respectively.

<!-- grade_id: cell-2b43d1c479894e01, points: 0, grade: false, solution: true -->
```python
grouped = df_breast_cancer.groupby('class')
num_benign = len(grouped.get_group('benign'))
num_malignant = len(grouped.get_group('malignant'))

```

<!-- grade_id: cell-f57e79bb105392be, points: 1, grade: true, solution: false -->
```python
print(f'Benign count: {num_benign}. Malignant count: {num_malignant}')

assert num_benign == 357
assert num_malignant == 212
```

Output:
```
Benign count: 357. Malignant count: 212
```

```python
# Store input and labels
x = dataset.data
y = dataset.target

print(f'x: {x.shape}. y: {y.shape}')
```

Output:
```
x: (569, 30). y: (569,)
```

### Question

Split the dataset into training and test sets with a 25% ratio. Use variables `x_train`, `y_train`, `x_test` and `y_test`.

<!-- grade_id: cell-a166a29ef1f6971f, points: 0, grade: false, solution: true -->
```python
x_train, x_test, y_train, y_test = train_test_split(x, y, test_size=.25)
```

<!-- grade_id: cell-e954af87bdf5fb27, points: 1, grade: true, solution: false -->
```python
print(f'x_train: {x_train.shape}. y_train: {y_train.shape}')
print(f'x_test: {x_test.shape}. y_test: {y_test.shape}')

assert x_train.shape == (426, 30)
assert y_train.shape == (426, )
assert x_test.shape == (143, 30)
assert y_test.shape == (143,)
```

Output:
```
x_train: (426, 30). y_train: (426,)
x_test: (143, 30). y_test: (143,)
```

### Question

Scale features by standardization while preventing information leakage from the test set.

<!-- grade_id: cell-5f0e1321a1045752, points: 0, grade: false, solution: true -->
```python
scaler = StandardScaler().fit(x_train)

x_train = scaler.transform(x_train)
x_test = scaler.transform(x_test)
```

<!-- grade_id: cell-63467a58075811c7, points: 1, grade: true, solution: false -->
```python
mean_train = x_train.mean()
std_train = x_train.std()
print(f'mean_train: {mean_train}. std_train: {std_train}')

assert np.abs(np.max(mean_train)) < 10**-6
assert np.abs(np.max(std_train - 1)) < 10**-6
```

Output:
```
mean_train: 2.80214036637786e-16. std_train: 1.0
```

## Step 3: Creating a classifier

### Question

Create a `KNeighborsClassifier` instance using only one nearest neighbor, store it into the `model` variable, and fit the training data.

<!-- grade_id: cell-6d03d4ed91b5b606, points: 0, grade: false, solution: true -->
```python
model = KNeighborsClassifier(n_neighbors=1)
model.fit(x_train, y_train)
```

Output:
```
KNeighborsClassifier(algorithm='auto', leaf_size=30, metric='minkowski',
                     metric_params=None, n_jobs=None, n_neighbors=1, p=2,
                     weights='uniform')
```

## Step 4: Evaluating the classifier

<!-- grade_id: cell-f9fdf200e5af12db, points: 1, grade: true, solution: false -->
```python
# Compute accuracy on training and test sets
train_acc = model.score(x_train, y_train)
test_acc = model.score(x_test, y_test)

print(f'Training accuracy: {train_acc * 100:.2f}%')
print(f'Test accuracy: {test_acc * 100:.2f}%')
```

Output:
```
Training accuracy: 100.00%
Test accuracy: 96.50%
```

### Question

Display precision, recall and f1-score for the classifier on test data. Interpret the results.

<!-- grade_id: cell-2fc9064207c0ad74, points: 0, grade: false, solution: true -->
```python
from sklearn import metrics
y_pred = model.predict(x_test)
precision = metrics.precision_score(y_test, y_pred, average=None)
print(f'Precision: {precision}')

# Using scikit-learn's recall_score() function
recall = metrics.recall_score(y_test, y_pred, average=None)
print(f'Recall: {recall}')
f1_score = metrics.f1_score(y_test, y_pred, average=None)
print(f'f1_score: {f1_score}')
```

Output:
```
Precision: [0.95555556 0.96938776]
Recall: [0.93478261 0.97938144]
f1_score: [0.94505495 0.97435897]
```

### Question

Go back to step 3 and try to find the best value for the `k` number of nearest neighbors.

```python
for i in range(8):
    model = KNeighborsClassifier(n_neighbors=1)
    model.fit(x_train, y_train)
    # Compute accuracy on training and test sets
    train_acc = model.score(x_train, y_train)
    test_acc = model.score(x_test, y_test)
    print(f'-------Data for k = {i}%-------')
    print(f'Training accuracy: {train_acc * 100:.2f}%')
    print(f'Test accuracy: {test_acc * 100:.2f}%')
    y_pred = model.predict(x_test)
    precision = metrics.precision_score(y_test, y_pred, average=None)
    print(f'Precision: {precision}')
    f1_score = metrics.f1_score(y_test, y_pred, average=None)
    print(f'f1_score: {f1_score}')
    
```

Output:
```
-------Data for k = 0%-------
Training accuracy: 100.00%
Test accuracy: 96.50%
Precision: [0.95555556 0.96938776]
f1_score: [0.94505495 0.97435897]
-------Data for k = 1%-------
Training accuracy: 100.00%
Test accuracy: 96.50%
Precision: [0.95555556 0.96938776]
f1_score: [0.94505495 0.97435897]
-------Data for k = 2%-------
Training accuracy: 100.00%
Test accuracy: 96.50%
Precision: [0.95555556 0.96938776]
f1_score: [0.94505495 0.97435897]
-------Data for k = 3%-------
Training accuracy: 100.00%
Test accuracy: 96.50%
Precision: [0.95555556 0.96938776]
f1_score: [0.94505495 0.97435897]
-------Data for k = 4%-------
Training accuracy: 100.00%
Test accuracy: 96.50%
Precision: [0.95555556 0.96938776]
f1_score: [0.94505495 0.97435897]
-------Data for k = 5%-------
Training accuracy: 100.00%
Test accuracy: 96.50%
Precision: [0.95555556 0.96938776]
f1_score: [0.94505495 0.97435897]
-------Data for k = 6%-------
Training accuracy: 100.00%
Test accuracy: 96.50%
Precision: [0.95555556 0.96938776
```