<!--NAVIGATION-->

<a href="https://colab.research.google.com/github/bpesquet/machine-learning-katas/blob/master/notebooks/katas/algorithms/RandomForest_Diabetes.ipynb"><img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open in Colab" title="Open in Google Colaboratory"></a>


## Instructions

This is a self-correcting exercise generated by [nbgrader](https://github.com/jupyter/nbgrader). 

Fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Run subsequent cells to check your code.

---

# Kata: Predict Diabetes Evolution

In this kata, you'll use several regression models to predict the disease progression one year after.

The Diabetes dataset contains ten baseline variables, age, sex, body mass index, average blood pressure, and six blood serum measurements were obtained for each of n = 442 diabetes patients, as well as the response of interest, a quantitative measure of disease progression one year after baseline.

## Package setup

```python
# Import base packages
import numpy as np
import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
import pandas as pd
```

### Question

Import the needed packages.

<!-- grade_id: cell-a3604c4afc82c828, points: 0, grade: false, solution: true -->
```python
# Import ML packages (edit this list if needed)
from sklearn.datasets import load_diabetes
from sklearn.tree import DecisionTreeClassifier, plot_tree
from sklearn.model_selection import train_test_split
from sklearn.metrics import mean_squared_error
```

## Step 1: Loading the data

```python
dataset = load_diabetes()

# Put data in a pandas DataFrame
df_diab = pd.DataFrame(dataset.data, columns=dataset.feature_names)
# Add target to DataFrame
df_diab['target'] = dataset.target
# Show 10 random samples
df_diab.sample(n=10)
```

Output:
```
age       sex       bmi        bp        s1        s2        s3  \
41  -0.099961 -0.044642 -0.067641 -0.108957 -0.074494 -0.072712  0.015505   
307  0.067136  0.050680 -0.030996  0.004658  0.024574  0.035638 -0.028674   
253  0.081666 -0.044642  0.033673  0.008101  0.052093  0.056619 -0.017629   
371  0.052606  0.050680 -0.009439  0.049415  0.050717 -0.019163 -0.013948   
225  0.030811  0.050680  0.032595  0.049415 -0.040096 -0.043589 -0.069172   
306  0.009016  0.050680 -0.001895  0.021872 -0.038720 -0.024800 -0.006584   
120 -0.049105 -0.044642  0.004572  0.011544 -0.037344 -0.018537 -0.017629   
382  0.048974 -0.044642  0.060618 -0.022885 -0.023584 -0.072712 -0.043401   
55  -0.041840 -0.044642 -0.049318 -0.036656 -0.007073 -0.022608  0.085456   
125 -0.005515  0.050680 -0.008362 -0.002228 -0.033216 -0.063630 -0.036038   

           s4        s5        s6  target  
41  -0.039493 -0.049868 -0.009362    55.0  
307  0.034309  0.023375  0.081764   172.0  
253  0.034309  0.034864  0.069
```

## Step 2: Preparing the data

### Question

Split the dataset into training and test sets with a 20% ratio.

<!-- grade_id: cell-4b94ad7a9c319486, points: 0, grade: false, solution: true -->
```python
x_train, x_test, y_train, y_test = train_test_split(dataset.data, dataset.target, test_size =0.2)
```

<!-- grade_id: cell-7135eb33d9bae222, points: 0, grade: true, solution: false -->
```python
print(f'x_train: {x_train.shape}. y_train: {y_train.shape}')
print(f'x_test: {x_test.shape}. y_test: {y_test.shape}')

assert x_train.shape == (353, 10)
assert y_train.shape == (353,)
assert x_test.shape == (89, 10)
assert y_test.shape == (89,)
```

Output:
```
x_train: (353, 10). y_train: (353,)
x_test: (89, 10). y_test: (89,)
```

## Step 3: Training several models

### Question

Create and train a `DecisionTreeRegressor`, a `GradientBoostingRegressor` and a `RandomForestRegressor` on the training data.

Compute their MSE on the training and test data.

<!-- grade_id: cell-f80bcaf317c3e6a7, points: 0, grade: false, solution: true -->
```python
from sklearn import tree
tree_reg = tree.DecisionTreeRegressor(random_state = 42)
tree_reg.fit(x_train, y_train)
```

Output:
```
DecisionTreeRegressor(criterion='mse', max_depth=None, max_features=None,
                      max_leaf_nodes=None, min_impurity_decrease=0.0,
                      min_impurity_split=None, min_samples_leaf=1,
                      min_samples_split=2, min_weight_fraction_leaf=0.0,
                      presort=False, random_state=42, splitter='best')
```

```python
train_acc = tree_reg.score(x_train, y_train)
test_acc = tree_reg.score(x_test, y_test)

print("Train accuracy", train_acc)
print("Test accuracy", test_acc)
```

Output:
```
Train accuracy 1.0
Test accuracy 0.15061418165837814
```

<!-- grade_id: cell-0ee3ffaa00916222, points: 0, grade: false, solution: true -->
```python
from sklearn.ensemble import GradientBoostingRegressor
grad_boost_reg = GradientBoostingRegressor(random_state=42)
grad_boost_reg.fit(x_train, y_train)
```

Output:
```
GradientBoostingRegressor(alpha=0.9, criterion='friedman_mse', init=None,
                          learning_rate=0.1, loss='ls', max_depth=3,
                          max_features=None, max_leaf_nodes=None,
                          min_impurity_decrease=0.0, min_impurity_split=None,
                          min_samples_leaf=1, min_samples_split=2,
                          min_weight_fraction_leaf=0.0, n_estimators=100,
                          n_iter_no_change=None, presort='auto',
                          random_state=42, subsample=1.0, tol=0.0001,
                          validation_fraction=0.1, verbose=0, warm_start=False)
```

```python
train_acc = grad_boost_reg.score(x_train, y_train)
test_acc = grad_boost_reg.score(x_test, y_test)

print("Train accuracy", train_acc)
print("Test accuracy", test_acc)
```

Output:
```
Train accuracy 0.8568189555005084
Test accuracy 0.40057156115115466
```

<!-- grade_id: cell-982fc1c1517d3914, points: 0, grade: false, solution: true -->
```python
from sklearn.ensemble import RandomForestRegressor
random_forest_reg = RandomForestRegressor(random_state=42)
random_forest_reg.fit(x_train, y_train)
```

Output:
```
/home/kimnovak/anaconda3/lib/python3.7/site-packages/sklearn/ensemble/forest.py:245: FutureWarning: The default value of n_estimators will change from 10 in version 0.20 to 100 in 0.22.
  "10 in version 0.20 to 100 in 0.22.", FutureWarning)
RandomForestRegressor(bootstrap=True, criterion='mse', max_depth=None,
                      max_features='auto', max_leaf_nodes=None,
                      min_impurity_decrease=0.0, min_impurity_split=None,
                      min_samples_leaf=1, min_samples_split=2,
                      min_weight_fraction_leaf=0.0, n_estimators=10,
                      n_jobs=None, oob_score=False, random_state=42, verbose=0,
                      warm_start=False)
```

<!-- grade_id: cell-7d3465c2a61dab34, points: 0, grade: false, solution: true -->
```python
train_acc = random_forest_reg.score(x_train, y_train)
test_acc = random_forest_reg.score(x_test, y_test)

print("Train accuracy", train_acc)
print("Test accuracy", test_acc)
```

Output:
```
Train accuracy 0.8937713185701133
Test accuracy 0.3843917870106892
```